```python
# Initialize Otter
import otter
grader = otter.Notebook("hw07.ipynb")
```

# Homework 7: Probability and Estimators
## Due Date: Thursday, April 6th, 11:59 PM PDT

You must submit this assignment to Gradescope by the on-time deadline, Thursday, April 6th, 11:59pm. Please read the syllabus for the grace period policy. No late
submissions beyond the grace period will be accepted. **We strongly encourage you to plan to submit your work to Gradescope several hours before the stated deadline.** This way, you will have ample time to reach out to staff for support if you encounter difficulties with submission. While course staff is happy to help guide you with submitting your assignment ahead of the deadline, we will not respond to last-minute requests for assistance (TAs need to sleep, after all!).

Please read the instructions carefully to submit your work to both the coding and written portals of Gradescope.


## Content Warning

This assignment includes an analysis of daily COVID-19 cases by U.S. county through 2021. If you feel uncomfortable with this topic, **please contact your GSI or the instructors.**



## Collaboration Policy

Data science is a collaborative activity. While you may talk with others about the homework, we ask that you **write your solutions individually**. If you do discuss the assignments with others please **include their names** below.

**Collaborators**: *list collaborators here*

## Introduction

In this homework, we will investigate a dataset that contains information about COVID-19 cases in the United States, vaccination rates, and various other metadata that can assist in modeling various aspects of COVID-19.

Through this homework assignment, you will demonstrate your experience with:
* Bootstrap sampling
* Bias-variance tradeoff and decomposition
* Multicollinearity in features

## Grading
Grading is broken down into autograded answers and free response. 

For autograded answers, the results of your code are compared to provided and/or hidden tests.

For free response, readers will evaluate how well you answered the question and/or fulfilled the requirements of the question. 

### Score breakdown

Question | Manual | Points
--- |---| ---
1a| No | 3
1b| No | 1
1c| Yes | 3
1d| Yes | 3
2a| No | 4
2b| Yes |3
2c| Yes | 1
3a| No | 2
3b| No | 4
3c| No | 2
3d| Yes | 3
4a| No | 3
4b| Yes | 3
4c| No | 2
4d| Yes | 2
Total | 7 | 39

```python
# Run this cell to set up your notebook
import numpy as np
import pandas as pd
import sklearn.linear_model as lm
from sklearn.model_selection import train_test_split

import matplotlib.pyplot as plt
import seaborn as sns
from IPython.display import Markdown

import scipy.stats

import warnings
warnings.filterwarnings("ignore")
```

<br/><br/><br/>

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Question 1: Exploratory Data Analysis

Let's perform some initial exploratory data analysis to examine and visualize potential trends in a COVID-19 dataset.

```python
# Run this cell to load the data, no further action needed.
covid_data = pd.read_csv('data/covid_data.csv')
covid_data.head(5)
```

The data are at county granularity; each row corresponds to COVID-19 data from one U.S. county. Here are some highlights and data sources:

* The first few columns encode county and state data; for example, check out the [Federal Information Processing System (FIPS)](https://transition.fcc.gov/oet/info/maps/census/fips/fips.txt) numeric encoding for U.S. counties.
* The next 600 columns record daily COVID-19 cases in the county for the date range 1/22/2020 to 9/12/2021. COVID-19 case data are from Center for Systems Sciences and Engineering (CSSE) at Johns Hopkins University [GitHub](https://github.com/CSSEGISandData/COVID-19/blob/master/csse_covid_19_data/csse_covid_19_time_series/time_series_covid19_confirmed_US.csv).
* The next few columns include county populations from [U.S. census data](https://www2.census.gov/programs-surveys/popest/datasets/2010-2020/counties/totals/co-est2020.csv), the latest of which is 2020.
* The last 5 columns record mask usage survey data on a 5-point scale from `NEVER` to `ALWAYS`. Data are collected in July 2020 from the New York Times [GitHub](https://github.com/nytimes/covid-19-data/blob/master/mask-use/mask-use-by-county.csv). Each column represents the proportion of population in that county who never/rarely/sometimes/frequently/always wear masks. Note, for a particular row, the numbers in those five columns sum up to $1$.

We can use `covid_data.describe()` to see various statistics about the numerical features of the provided COVID-19 data. Do any particular statistics stand out to you? Which might be useful when modeling?

**Note:** This isn't a question (i.e. it's worth no points); This is just food for thought as you start to explore the dataset.

```python
# Run this cell to load the data see see statistics of the covid_data
covid_data.describe()
```

<br><br>

---

### Question 1a

In this homework, we will use linear regression to predict **the number of COVID-19 cases  per capita on September 12th, 2021**. Define a column `'9/12/2021_cpc'` in `covid_data` corresponding to the number of cases per capita on September 12th, 2021. 

Note that we will **always** use the `'POPESTIMATE2020'` as the population of each county.

*Hint*: The number of cases per capita should be the total number of cases in a county divided by the population of the county.


```python
...
```

```python
grader.check("q1a")
```

<br><br>

---

### Question 1b

Assign `mask_data` that has six columns from the original `covid_data` table: the five mask usage columns described earlier and the `9/12/2021_cpc` column.

**Note**: You should make a **copy** of these columns using `df.copy()` ([link](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.copy.html)).

```python
mask_data = ...
mask_data
```

```python
grader.check("q1b")
```

<!-- BEGIN QUESTION -->

<br><br>

---

### Question 1c

Our goal is to use county-wise mask usage data to predict the number of COVID-19 cases per capita on September 12th, 2021 (i.e., the column `9/12/2021_cpc`). But before modeling, let's do some EDA to explore the multicollinearality in these features, and then we will revisit this question in part 4. 

Create a visualization that shows the pairwise correlation between each combination of columns in `mask_data`. For 2-D visualizations, consider Seaborn's [heatmap](https://seaborn.pydata.org/generated/seaborn.heatmap.html). Remember to add a title to your plot.

**Hint**: You should be plotting 36 values corresponding to the [pairwise correlations](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.corr.html) of the six columns in `mask_data`.


```python
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<br><br>

---

### Question 1d

(1) Describe the trends and takeaways visible in the visualization of pairwise correlations you plotted in Question 1c. Specifically, how does the correlation between pairs of features (i.e. mask usage) look like? How does the correlation between mask usage and cases per capita look like?

(2) If we are going to build a linear regression model (with an intercept term) using all five mask usage columns as features, then what problem will we encounter?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br/><br/><br/>

<hr style="border: 1px solid #fdb515;" />

## Question 2: Creating a Preliminary COVID-19 Model

This question will guide you through creating a linear regression model that will predict the number of COVID-19 cases per capita given various COVID-19 safety protocols that have been implemented. Then, we will investigate the bias, variance, and observational noise of this model in the next two questions.


<br><br>

---

### Question 2a

Despite the problems we see previously, let's train a linear regression model with an intercept term, using Scikit-learn, to predict the number of COVID-19 cases per capita for September 12, 2021 using county-wise mask usage data from `mask_data`. Use `train_test_split` to evaluate your model's RMSE on a held-out test set with 33% of the COVID-19 data; call the resulting splits `X_train`, `X_test`, `y_train`, and `y_test`.

To pass the autograder, make sure to set the parameter `random_state` to 42 in your call to `train_test_split` to generate a reproducible data split ([documentation](http://scikit-learn.org/stable/modules/generated/sklearn.model_selection.train_test_split.html)).



```python
# Create train/test sets
X = ...
y = ...
X_train, X_test, y_train, y_test = ...

# Fit the linear model and make predictions (you will need multiple lines)
...

# Compute RMSE on train and test sets
train_rmse_cpc = ...
test_rmse_cpc = ...


train_rmse_cpc, test_rmse_cpc
```

```python
grader.check("q2a")
```

<!-- BEGIN QUESTION -->

<br><br>

---

### Question 2b

To visualize the model performance from part (a), let's make the following two visualizations: 

(1) The predicted values vs. observed values on the test set,

(2) The residuals plot. (Note: in multiple linear regression, the residual plot has predicted values vs. residuals) 

**Note:**
* We've used `plt.subplot` ([documentation](https://matplotlib.org/stable/gallery/subplots_axes_and_figures/subplots_demo.html)) so that you can view both visualizations side-by-side. For example, `plt.subplot(121)` sets the plottable area to the first column of a 1x2 plot grid; you can then call Matplotlib and Seaborn functions to plot that area, before the next `plt.subplot(122)` area is set.
* Remember to add a guiding line to both plot where $\hat{y} = y$, i.e., where the residual is 0.
* Please add descriptive titles and axis labels for your plots!

```python
plt.figure(figsize=(12,6))      # do not change this line
plt.subplot(121)                # do not change this line
# (1) predictions vs. observations
...

plt.subplot(122)               # do not change this line
# (2) residual plot
...

plt.tight_layout()             # do not change this line
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<br><br>

---

### Question 2c

Describe what the plots in part (b) indicate about this linear model.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br/><br/><br/>

<hr style="border: 1px solid #fdb515;" />

## Question 3: Performing Multicollinearity Analysis

This question will guide you through performing an analysis that can reveal potential multicollinearity in our features, which is not ideal. In particular, we will use bootstrap to get $95\%$ confidence intervals on the fitted parameters. Here's a reminder of the outline of boostrapping:

1. (3a) Assume the sample is a representative sample, we can simulate different samples by resampling from the original sample.
2. (3b) Calculate statistic(s) for each resample to simulate different potential result. Statistic can be a mean (Lab 9), correlation (Lab 9), coefficients of linear model (Q3 below), and any function of the sample data.
3. (3c) Use the distribution of the above statistics calculated on each of the resample to construct confidence interval.
4. (3d) Use the result of confidence interval to make inference of the population.

<br><br>

---

### Question 3a

Fill in the blanks below to implement the `bootstrap_sample` function, that returns $k$ randomly drawn samples from a dataset `data` of size $n$ with replacement, each of size $n$ (i.e. same size as the dataset). In other words, the returned object should be a Python list `samples` containing $k$ Pandas DataFrames, each of which have $n$ rows.

**Hint**: You may find `df.sample` helpful, see [documentation](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.sample.html). This is very similar to Lab 8.


```python
def bootstrap_sample(data, k):
    """
    Performs bootstrap sampling on data to obtain k samples of size n.
    
    Arguments:
        data - Dataset contained as a Pandas DataFrame 
        k - Number of randomly drawn samples
    
    Returns:
        samples - List containing k Pandas DataFrames of size n each
                  corresponding to each sample  
    """
    samples = []
    ...

# Print out the first dataframe only
bootstrap_sample(mask_data, 1)[0]
```

```python
grader.check("q3a")
```

<br><br>

---

### Question 3b

Using the function from the previous part, let's do the following:

1. Generate 1000 bootstrapped samples from the original `mask_data` dataframe. 
2. Use Scikit-learn to fit a linear regression model (with an intercept term) to use the mask usage features to predict the `9/12/2021_cpc` response. 
3. Store each of the 1000 trained models in the Python list `models`.

**Hint:**
* You *should not* create any validation or testing sets in this subpart, your model should fit to the entire resampled dataframe.
* `LinearRegression` is an object type, to store a new model you must create a new instance first!

```python
np.random.seed(42)

datasets = ...
models = []
...


# Datasets take up a lot of memory, so we should remove them!
del datasets
```

```python
grader.check("q3b")
```

<br><br>

---

### Question 3c

Fill in the blanks below in the `confidence_interval` function to generate a $95\%$ confidence interval for each of our parameters $\theta_i$, including the intercept term $\theta_0$. All of the helper code to extract coefficients from our trained models has been implemented for you already.

**Hint**: 
- For a refresh on confidence intervals, refer to this link from the [Data 8 textbook](https://inferentialthinking.com/chapters/13/3/Confidence_Intervals.html). 
- Pay close attention to how the arrays used below are formated. What does each row represent? What does each column represent? To get the $i$th column from a 2D-array, you can use `2D_array[:, i]`.


```python
def extract_coefs(models, include_intercept = True):
    """
    NOTE: This function has already been implemented. You do not need to modify this!
    
    Extracts coefficients of all the linear regression models in MODELS, and returns
    it as a NumPy array with one model's coefficients as each row.
    
    Arguments:
        models - Contains k sklearn LinearRegression models, each with p + 1 coefficients
        include_intercept - Whether to include intercept in returned coefficients
    
    Returns:
        coef_array - Coefficients of all k models, each with p + 1 coefficients (if intercept
                     enabled, otherwise p). Returned object is k x (p + 1) NumPy array.
    """
    coef_array = np.zeros(shape = (len(models), len(models[0].coef_) + 1))
    for i, m in enumerate(models):
        coef_array[i, 0] = m.intercept_
        coef_array[i, 1:] = m.coef_
    if include_intercept:
        return coef_array 
    return coef_array[:, 1:]

def confidence_interval(coefs):
    """
    Calculates confidence intervals for each theta_i based on coefficients of 
    bootstrapped models. Returns output as a list of confidence intervals.
    
    Arguments:
        coefs - Output of extract_coefs, a k x (p + 1) or k x p NumPy array containing
                coefficients of bootstrapped models
    
    Returns:
        cis - Confidence intervals of each parameter theta_i in the form of a 
              list like this: [(0.5, 0.75), (0.2, 0.4), ...]
    """
    cis = []
    
    # FILL IN CODE BELOW
    for i in range(...):
        theta_i_values = ...
        theta_i_lower_ci, theta_i_upper_ci = np.percentile(...), np.percentile(...)
        cis.append((theta_i_lower_ci, theta_i_upper_ci))
    
    return cis


# Compute confidence intervals
model_coefs = extract_coefs(models)
cis = confidence_interval(model_coefs)

# Pretty print in a table
display(Markdown('#### Confidence Intervals:'))
md_list = ["|parameter|lower|upper|",
           "----|----|----|"]
md_list += [fr"|$\theta_{i}$|{lci}|{uci}|" for i, (lci, uci) in enumerate(cis)]
display(Markdown('\n'.join(md_list)))
```

```python
grader.check("q3c")
```

<!-- BEGIN QUESTION -->

<br><br>

---

### Question 3d

Interpret the confidence intervals above for each of the $\theta_i$, where $\theta_0$ is the intercept term and the remaining $\theta_i$'s are parameters corresponding to mask usage features. What does this indicate about our data and our model?

Describe a reason why this could be happening.

**Hint**: Take a look at the design matrix, heatmap, and response from Question 1!


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br/><br/><br/>

<hr style="border: 1px solid #fdb515;" />

## Question 4: Performing Bias-Variance Analysis

This question will guide you through performing an analysis that can estimate the bias and variance of our models, which can be helpful in modeling.

Recall that the **model variance** on a data point $\vec{x_k}$ is simply the variance of our predictions on that sample point $\vec{x_k}$. 

$$\text{model variance} = \mathrm{Var}(\hat{Y}(\vec{x_k}) = \mathrm{Var}(\vec{x_k}^T\hat{\theta})$$

Since we are using OLS, we can also rewrite our model as $\hat{Y}(\vec{x_k}) = \vec{x_k}^T\hat{\theta} = \hat{\theta_0}\times 1 + \hat{\theta_1}x_1 + ... + \hat{\theta_p}x_p$. To estimate the model variance, we can sample a particular data point $(\vec{x_k}, y_k)$, and calculate the variance of the predictions using different fitted model in `models`. Note that the varaince is taken acorss different models, $\vec{x_k}$ are pre-selected and fixed. 

Also recall that **model risk** for this point is the same as mean square error over all possible fitted models:

$$\text{model risk} = \mathbb{E}\left[\left(y_k - \hat{Y}(\vec{x_k} \right)^2\right] = \frac{1}{\text{# bootstrap}}\sum_{j=1}^{\text{# bootstrap}} (y_k - \hat{Y^j}(\vec{x_k} ))^2 = MSE(\vec{x_k})$$

Where the superscript is used to denote different models. Here, we are considering a particular observation of the random response variable $y_k$. Therefore model risk is an expectation over the estimate $\hat{\theta^j}$, coefficients corresponding to different models. $\hat{\theta^j}$ itself is a random varaible, because it was derived by fitting to a different random resample.


We can also find the **ratio of model variance to model risk**. You can interpret this ratio as the the proportion of the expected square error on the data point "captured"/"explained" by the model variance. 

$$
\frac{\text{model variacne}}{\text{model risk}}=\frac{\mathrm{Var}(\hat{Y}(\vec{x_k}))}{\mathbb{E}\left[\left(y_k - \hat{Y}(\vec{x_k} \right)^2\right]}
$$

<br><br>

---

### Question 4a

Let's use the 100th data point $(\vec{x_{100}}, y_{100})$, i.e. k =100 to find these quantities!

Complete the function `simulate` that takes in a data point `xk`, `yk` and return `model_risk`, `model_var`, and `ratio`, ratio of model variance to model risk.

Here is a suggested format but you do not need to follow the skeleton code:
* Assign `predictions` to a length 1000 vector where each element is the prediction on `xk` using a model in `models`. Note that sklearn's `.predict` return an array, but we only need a scaler prediction! Try to print out a prediction and ajust your code if needed.
* Use `predictions` to compute `model_risk`.
* Use `predictions` to compute `model_var`.
* Use `model_risk` and `model_var` to compute the `ratio`, ratio of model variance to model risk.

**Hint**: Use list comprehension for creating your predictions to save overhead of using multiple appends.

```python
def simulate(xk, yk, models):
    predictions = ...
    model_risk = ...
    model_var = ...
    ratio = ...
    return model_risk, model_var, ratio
...
model_risk, model_var, ratio = simulate(x_100, y_100, models)
model_risk, model_var, ratio
```

```python
grader.check("q4a")
```

<!-- BEGIN QUESTION -->

<br><br>

---

### Question 4b

Comment on the ratio `ratio`, which is the proportion of the expected square error on the data point captured by the model variance. Is the model variance the dominant term in the bias-variance decomposition? If not, what term(s) dominate the bias-variance decomposition?

**Note**: The Bias-Variance decomposition from lecture:

$$
\text{model risk} = \sigma^2 + (\text{model bias})^2 + \text{model variance}
$$

where $\sigma^2$ is the observation variance, or "irreducible error".


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br><br>

---

### Question 4c

Now let's calculate the average variance and average mean square error across 250 randomly sampled $(x_i, y_i)$ points. In other words, estimate the following quantities across all $x_i$ and $y_i$ in `X_sample` and `y_sample`:

$$
\frac{1}{250} \sum_{k=1}^{250} \mathrm{Var}\left(\hat{Y}(\vec{x_k})\right)
$$

and

$$
\frac{1}{250} \sum_{k=1}^{250} \mathbb{E}\left[ (y_k - \hat{Y}(\vec{x_k}))^2 \right]
$$

**Hint:** Call `simulate` from 4a.


```python
np.random.seed(42)

X_sample = X.sample(250)         # Generate 250 x_i
y_sample = y.loc[X_sample.index] # ...and select the corresponding y_i

var, mse = [], []
...

avg_var, avg_mse = np.mean(var), np.mean(mse)
avg_var, avg_mse
```

```python
grader.check("q4c")
```

<!-- BEGIN QUESTION -->

<br><br>

---

### Question 4d

Propose a solution to reducing the mean square error using the insights gained from the bias-variance decomposition above.

Assume that the standard bias-variance decomposition used in lecture can be applied here.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br><br>


<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Congratulations! You have finished Homework 7!

Here are some cats :) Can you find their names? There are two ways you can find it on JupyterLab.

<img src="Pandas.jpg" width="370px"/> <img src="Patches.jpg" width="200px" /> <img src="Pishi.jpg" width="208px" />


Below, you will see two cells. Running the first cell will automatically generate a PDF of all questions that need to be manually graded, and running the second cell will automatically generate a zip with your autograded answers. **You are responsible for both the coding portion (the zip from Homework 7) and the written portion (the PDF with from Homework 7) to their respective Gradescope portals.** The coding proportion should be submitted to Homework 7 Coding as a single zip file, and the written portion should be submitted to Homework 7 Written as a single pdf file. When submitting the written portion, please ensure you select pages appropriately and check all plots appear. 

If there are issues with automatically generating the PDF in the first cell, you can try downloading the notebook as a PDF by clicking on `File -> Save and Export Notebook As... -> PDF`. If that doesn't work either, you can manually take screenshots of your answers to the manually graded questions and submit those. Either way, **you are responsible for ensuring your submission follows our requirements, we will NOT be granting regrade requests for submissions that don't follow instructions.**


```python
from otter.export import export_notebook
from os import path
from IPython.display import display, HTML
export_notebook("hw07.ipynb", filtering=True, pagebreaks=True)
if(path.exists('hw07.pdf')):
    display(HTML("Download your PDF <a href='hw07.pdf' download>here</a>."))
else:
    print("\n Pdf generation fails, please try the other methods described above")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1a (points: 3)
```
>>> np.isclose(covid_data['9/12/2021_cpc'].iloc[0], 0.16541098940244012)
True
```
```
>>> np.isclose(covid_data['9/12/2021_cpc'].sum(), 397.68435739730006)
True
```


Tests for q1b (points: 1)
```
>>> mask_data.shape == (3141, 6) or mask_data.shape == (3141, 7)
True
```


Tests for q2a (points: 4)
```
>>> 0.01 <= train_rmse_cpc <= 0.05
True
```
```
>>> 0.01 <= test_rmse_cpc <= 0.05
True
```


Tests for q3a (points: 2)
```
>>> all([len(df_i) == len(mask_data) for df_i in bootstrap_sample(mask_data, 1)])
True
```
```
>>> len(bootstrap_sample(mask_data, 3)) == 3
True
```
```
>>> np.isclose(bootstrap_sample(mask_data, 1)[0]['NEVER'].mean(), 0.0799, atol = 8e-3)
True
```


Tests for q3b (points: 4)
```
>>> len(models) == 1000
True
```
```
>>> all([isinstance(model, lm.LinearRegression) for model in models])
True
```
```
>>> all([len(model.coef_) == 5 and model.intercept_ != 0 for model in models])
True
```
```
>>> all(models[0].coef_ != models[1].coef_)
True
```


Tests for q3c (points: 2)
```
>>> len(cis) == 6 and all(len(ci) == 2 for ci in cis)
True
```
```
>>> all([lci < 0 and uci > 0 for lci, uci in cis])
True
```


Tests for q4a (points: 3)
```
>>> model_risk <= 0.003
True
```
```
>>> model_var <= 1e-5
True
```
```
>>> np.isclose(ratio, 0.001316917649910885, atol = 1e-3)
True
```


Tests for q4c (points: 2)
```
>>> avg_var <= 1e-5
True
```
```
>>> avg_mse <= 5e-3
True
```
